<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<question type="coderunner">
    <name>
      <text>Question Time</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p dir="ltr" style="text-align: left;"></p>
<h3> </h3>
<h3>Exercise 2: Question Time</h3>
<p>Write a simple C program that prompts the user a series of questions:</p>
<p></p>
<ul>
<li>What is your name?</li>
<li>How old are you?</li>
<li>What is your favourite animal?</li>
</ul>
<p></p>
<p>The answers should be safely stored in three variables and then printed to the screen in the form of a sentence that includes each of the answers.</p>
<p></p>
<p>Save your code in a file called <strong>exercise02.c</strong></p>
<p>You can use any text editor you like to create the file, but it must compile and run using the following commands (on the terminal of a University Linux machine):</p>
<pre>gcc exercise02.c -pedantic -Wall -o exercise02<br>./exercise02</pre>
<p>Attempt the quiz when your code is working to receive your mark</p>
<p>If you think this quiz is behaving incorrectly, speak to Andrew or email him (andrew.fagan@strath.ac.uk)</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>c_program</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[#include <stdio.h>

int main(){
	char name[11];
	char age[4];
	char animal[21];
	
	printf("What is your name?\n");
	scanf(" %10s", name);
	
	printf("How old are you?\n");
	scanf(" %3s", age);
	
	printf("What is your favourite animal?\n");
	scanf(" %20s", animal);
	
	
	
	printf("You are %s, you are %s years old and your favourite animal is %s", name, age, animal);
}]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader>RegexGrader</grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>Clemens
180
bats</text>
      </stdin>
      <expected>
                <text>(?=.*Clemens)(?=.*bats)(?=.*180)</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>Fred
40
cats</text>
      </stdin>
      <expected>
                <text>(?=.*Fred)(?=.*cats)(?=.*40)</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text></text>
      </testcode>
      <stdin>
                <text>Andrew
28
dogs</text>
      </stdin>
      <expected>
                <text>(?=.*Andrew)(?=.*dogs)(?=.*28)</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>
